unsat
(assume a0 (distinct y$n1s2 y$n0s2 y$n2s2 y$n3s2))
(assume a1 (distinct y$n0s3 y$n6s3 y$n7s3 y$n1s3 y$n4s3 y$n5s3 y$n2s3 y$n3s3))
(assume a2 (distinct y$n0s5 y$n1s5 y$n31s5))
(assume a3 (not (= y$n0s8 y$n128s8)))
(assume a4 (= y$2 (= y$n0s3 y$i_rx_phy.bit_cnt)))
(assume a5 (= y$5 (= y$n1s2 y$i_rx_phy.dpll_state)))
(assume a6 (= y$i_rx_phy.fs_ce (not y$7)))
(assume a7 (= y$i_rx_phy.fs_ce_d (not y$9)))
(assume a8 (= y$i_rx_phy.fs_ce_r1 (not y$11)))
(assume a9 (= y$i_rx_phy.fs_ce_r2 (not y$13)))
(assume a10 (= y$i_rx_phy.fs_ce_r3 (not y$15)))
(assume a11 (= y$17 (= y$n0s3 y$i_rx_phy.fs_state)))
(assume a12 (= y$20 (= y$n0s8 y$i_rx_phy.hold_reg)))
(assume a13 (= y$22 (= y$n0s3 y$i_rx_phy.one_cnt)))
(assume a14 (= y$i_rx_phy.rx_active (! (not y$24) :named @p_1)))
(assume a15 (= y$i_rx_phy.rx_en (not y$26)))
(assume a16 (= y$i_rx_phy.rx_valid (! (not y$28) :named @p_2)))
(assume a17 (= y$i_rx_phy.rx_valid1 (not y$30)))
(assume a18 (= y$i_rx_phy.rx_valid_r (not y$32)))
(assume a19 (= y$i_rx_phy.rxd_s (not y$34)))
(assume a20 (= y$i_rx_phy.rxd_s1 (not y$36)))
(assume a21 (= y$i_rx_phy.rxd_t1 (not y$38)))
(assume a22 (= y$i_rx_phy.rxdn_s (not y$40)))
(assume a23 (= y$i_rx_phy.rxdn_s1 (not y$42)))
(assume a24 (= y$i_rx_phy.rxdn_s1r (not y$44)))
(assume a25 (= y$i_rx_phy.rxdn_t1 (not y$46)))
(assume a26 (= y$i_rx_phy.rxdp_s (not y$48)))
(assume a27 (= y$i_rx_phy.rxdp_s1 (not y$50)))
(assume a28 (= y$i_rx_phy.rxdp_s1r (not y$52)))
(assume a29 (= y$i_rx_phy.rxdp_t1 (not y$54)))
(assume a30 (= y$i_rx_phy.sd_nrzi (not y$56)))
(assume a31 (= y$i_rx_phy.sd_r (not y$58)))
(assume a32 (= y$i_rx_phy.shift_en (not y$60)))
(assume a33 (= y$i_rx_phy.synced_d (not y$62)))
(assume a34 (= y$i_tx_phy.TxReady_o (not y$64)))
(assume a35 (= y$i_tx_phy.append_eop (not y$66)))
(assume a36 (= y$i_tx_phy.append_eop_sync1 (not y$68)))
(assume a37 (= y$i_tx_phy.append_eop_sync2 (not y$70)))
(assume a38 (= y$i_tx_phy.append_eop_sync3 (not y$72)))
(assume a39 (= y$74 (= y$n0s3 y$i_tx_phy.bit_cnt)))
(assume a40 (= y$i_tx_phy.data_done (not y$76)))
(assume a41 (= y$78 (= y$n0s8 y$i_tx_phy.hold_reg)))
(assume a42 (= y$i_tx_phy.ld_data (not y$80)))
(assume a43 (= y$i_tx_phy.ld_data_d (not y$82)))
(assume a44 (= y$i_tx_phy.ld_eop_d (not y$84)))
(assume a45 (= y$i_tx_phy.ld_sop_d (not y$86)))
(assume a46 (= y$88 (= y$n0s3 y$i_tx_phy.one_cnt)))
(assume a47 (= y$i_tx_phy.sd_bs_o (not y$90)))
(assume a48 (= y$i_tx_phy.sd_raw_o (not y$93)))
(assume a49 (= y$i_tx_phy.sft_done (not y$95)))
(assume a50 (= y$i_tx_phy.sft_done_r (not y$97)))
(assume a51 (= y$99 (= y$n0s3 y$i_tx_phy.state)))
(assume a52 (= y$i_tx_phy.tx_ip (not y$101)))
(assume a53 (= y$i_tx_phy.tx_ip_sync (not y$103)))
(assume a54 (= y$i_tx_phy.tx_ready_d (not y$105)))
(assume a55 (= y$i_tx_phy.txdn (not y$107)))
(assume a56 (= y$i_tx_phy.txoe_r1 (not y$111)))
(assume a57 (= y$i_tx_phy.txoe_r2 (not y$113)))
(assume a58 (= y$116 (= y$n0s5 y$rst_cnt)))
(assume a59 (= y$usb_rst (not y$118)))
(assume a60 (= y$prop (not y$967)))
(assume a61 (! (= y$w$1_op (Concat_32_1_31 y$i_rx_phy.rx_valid y$n0s31)) :named @p_3))
(assume a62 (! (= y$941 (not (= y$w$1_op y$n1s32))) :named @p_4))
(assume a63 (! (= y$w$2_op (Concat_32_1_31 y$i_rx_phy.rx_active y$n0s31)) :named @p_5))
(assume a64 (! (= y$944 (= y$n1s32 y$w$2_op)) :named @p_6))
(assume a65 (= y$prop0_op (or y$941 y$944)))
(assume a66 (! (= y$947 (! (= y$prop y$prop0_op) :named @p_7)) :named @p_8))
(assume a67 (! (= y$978 (! (and y$2 y$5 y$7 y$9 y$11 y$13 y$15 y$17 y$20 y$22 y$24 y$26 y$28 y$30 y$32 y$34 y$36 y$38 y$40 y$42 y$44 y$46 y$48 y$50 y$52 y$54 y$56 y$58 y$60 y$62 y$64 y$66 y$68 y$70 y$72 y$74 y$76 y$78 y$80 y$82 y$84 y$86 y$88 y$90 y$i_tx_phy.sd_nrzi_o y$93 y$95 y$97 y$99 y$101 y$103 y$105 y$107 y$i_tx_phy.txdp y$i_tx_phy.txoe y$111 y$113 y$116 y$118 y$967 y$947) :named @p_9)) :named @p_10))
(assume a68 y$978)
(step t0 (cl (not (! (= (! (not y$944) :named @p_11) false) :named @p_12)) (not @p_11) false) :rule equiv_pos2)
(step t1 (cl (not (! (= @p_6 (! (= y$944 (! (= y$n1s32 (! (Concat_32_1_31 @p_1 y$n0s31) :named @p_19)) :named @p_20)) :named @p_42)) :named @p_43)) (not @p_6) @p_42) :rule equiv_pos2)
(step t2 (cl (= y$944 y$944)) :rule refl)
(step t3 (cl (= y$n1s32 y$n1s32)) :rule refl)
(step t4 (cl (not (! (= @p_5 (! (= y$w$2_op @p_19) :named @p_44)) :named @p_45)) (not @p_5) @p_44) :rule equiv_pos2)
(step t5 (cl (= y$w$2_op y$w$2_op)) :rule refl)
(step t6 (cl (= y$n0s31 y$n0s31)) :rule refl)
(step t7 (cl (= (Concat_32_1_31 y$i_rx_phy.rx_active y$n0s31) @p_19)) :rule cong :premises (a14 t6))
(step t8 (cl @p_45) :rule cong :premises (t5 t7))
(step t9 (cl @p_44) :rule resolution :premises (t4 t8 a63) :args (@p_45 false @p_5 false))
(step t10 (cl (= (= y$n1s32 y$w$2_op) @p_20)) :rule cong :premises (t3 t9))
(step t11 (cl @p_43) :rule cong :premises (t2 t10))
(step t12 (cl @p_42) :rule resolution :premises (t1 t11 a64) :args (@p_43 false @p_6 false))
(step t13 (cl (not (! (= (! (not y$941) :named @p_24) (! (= y$n1s32 (! (Concat_32_1_31 false y$n0s31) :named @p_15)) :named @p_23)) :named @p_25)) (not @p_24) @p_23) :rule equiv_pos2)
(step t14 (cl (not (! (= (! (= y$941 (! (not (! (= y$n1s32 y$w$1_op) :named @p_35)) :named @p_36)) :named @p_37) (! (= y$941 (! (not (! (= y$n1s32 (! (Concat_32_1_31 @p_2 y$n0s31) :named @p_29)) :named @p_30)) :named @p_31)) :named @p_34)) :named @p_38)) (not @p_37) @p_34) :rule equiv_pos2)
(step t15 (cl (= y$941 y$941)) :rule refl)
(step t16 (cl (not (! (= @p_3 (! (= y$w$1_op @p_29) :named @p_40)) :named @p_41)) (not @p_3) @p_40) :rule equiv_pos2)
(step t17 (cl (= y$w$1_op y$w$1_op)) :rule refl)
(step t18 (cl (= (Concat_32_1_31 y$i_rx_phy.rx_valid y$n0s31) @p_29)) :rule cong :premises (a16 t6))
(step t19 (cl @p_41) :rule cong :premises (t17 t18))
(step t20 (cl @p_40) :rule resolution :premises (t16 t19 a61) :args (@p_41 false @p_3 false))
(step t21 (cl (= @p_35 @p_30)) :rule cong :premises (t3 t20))
(step t22 (cl (= @p_36 @p_31)) :rule cong :premises (t21))
(step t23 (cl @p_38) :rule cong :premises (t15 t22))
(step t24 (cl (not (! (= @p_4 @p_37) :named @p_39)) (not @p_4) @p_37) :rule equiv_pos2)
(step t25 (cl (= (= y$w$1_op y$n1s32) @p_35)) :rule rare_rewrite :args ("eq-symm" y$w$1_op y$n1s32))
(step t26 (cl (= (not (= y$w$1_op y$n1s32)) @p_36)) :rule cong :premises (t25))
(step t27 (cl @p_39) :rule cong :premises (t15 t26))
(step t28 (cl @p_37) :rule resolution :premises (t24 t27 a62) :args (@p_39 false @p_4 false))
(step t29 (cl @p_34) :rule resolution :premises (t14 t23 t28) :args (@p_38 false @p_37 false))
(step t30 (cl (not (! (= y$28 (! (= y$28 true) :named @p_32)) :named @p_33)) @p_2 @p_32) :rule equiv_pos2)
(step t31 (cl (= @p_32 y$28)) :rule rare_rewrite :args ("bool-eq-true" y$28))
(step t32 (cl @p_33) :rule symm :premises (t31))
(step t33 (cl (not @p_10) (not y$978) @p_9) :rule equiv_pos2)
(step t34 (cl @p_9) :rule resolution :premises (t33 a67 a68) :args (@p_10 false y$978 false))
(step t35 (cl y$28) :rule and :premises (t34) :args (12))
(step t36 (cl @p_32) :rule resolution :premises (t30 t32 t35) :args (@p_33 false y$28 false))
(step t37 (cl (= @p_2 (! (not true) :named @p_13))) :rule cong :premises (t36))
(step t38 (cl (= @p_29 (! (Concat_32_1_31 @p_13 y$n0s31) :named @p_14))) :rule cong :premises (t37 t6))
(step t39 (cl (= @p_30 (! (= y$n1s32 @p_14) :named @p_26))) :rule cong :premises (t3 t38))
(step t40 (cl (= @p_31 (! (not @p_26) :named @p_27))) :rule cong :premises (t39))
(step t41 (cl (= y$941 @p_27)) :rule trans :premises (t29 t40))
(step t42 (cl (= @p_24 (! (not @p_27) :named @p_28))) :rule cong :premises (t41))
(step t43 (cl (= @p_28 @p_26)) :rule rare_rewrite :args ("bool-double-not-elim" @p_26))
(step t44 (cl (= @p_13 false)) :rule rare_rewrite :args ("evaluate"))
(step t45 (cl (= @p_14 @p_15)) :rule cong :premises (t44 t6))
(step t46 (cl (= @p_26 @p_23)) :rule cong :premises (t3 t45))
(step t47 (cl (= @p_28 @p_23)) :rule trans :premises (t43 t46))
(step t48 (cl @p_25) :rule trans :premises (t42 t47))
(step t49 (cl y$prop0_op (not (or y$941 y$944))) :rule equiv2 :premises (a65))
(step t50 (cl (not @p_8) (not y$947) @p_7) :rule equiv_pos2)
(step t51 (cl y$947) :rule and :premises (t34) :args (60))
(step t52 (cl @p_7) :rule resolution :premises (t50 a66 t51) :args (@p_8 false y$947 false))
(step t53 (cl y$prop (not y$prop0_op)) :rule equiv2 :premises (t52))
(step t54 (cl (not y$prop) (not y$967)) :rule equiv1 :premises (a60))
(step t55 (cl y$967) :rule and :premises (t34) :args (59))
(step t56 (cl (not y$prop)) :rule resolution :premises (t54 t55) :args (y$967 false))
(step t57 (cl (not y$prop0_op)) :rule resolution :premises (t53 t56) :args (y$prop true))
(step t58 (cl (not (or y$941 y$944))) :rule resolution :premises (t49 t57) :args (y$prop0_op true))
(step t59 (cl @p_24) :rule not_or :premises (t58) :args (0))
(step t60 (cl @p_23) :rule resolution :premises (t13 t48 t59) :args (@p_25 false @p_24 false))
(step t61 (cl (not (! (= y$24 (! (= y$24 true) :named @p_21)) :named @p_22)) @p_1 @p_21) :rule equiv_pos2)
(step t62 (cl (= @p_21 y$24)) :rule rare_rewrite :args ("bool-eq-true" y$24))
(step t63 (cl @p_22) :rule symm :premises (t62))
(step t64 (cl y$24) :rule and :premises (t34) :args (10))
(step t65 (cl @p_21) :rule resolution :premises (t61 t63 t64) :args (@p_22 false y$24 false))
(step t66 (cl (= @p_1 @p_13)) :rule cong :premises (t65))
(step t67 (cl (= @p_19 @p_14)) :rule cong :premises (t66 t6))
(step t68 (cl (= @p_20 (! (= @p_15 @p_14) :named @p_16))) :rule cong :premises (t60 t67))
(step t69 (cl (= y$944 @p_16)) :rule trans :premises (t12 t68))
(step t70 (cl (= @p_11 (! (not @p_16) :named @p_17))) :rule cong :premises (t69))
(step t71 (cl (! (= @p_15 @p_15) :named @p_18)) :rule refl)
(step t72 (cl (= @p_16 @p_18)) :rule cong :premises (t71 t45))
(step t73 (cl (= @p_18 true)) :rule rare_rewrite :args ("eq-refl" @p_15))
(step t74 (cl (= @p_16 true)) :rule trans :premises (t72 t73))
(step t75 (cl (= @p_17 @p_13)) :rule cong :premises (t74))
(step t76 (cl (= @p_17 false)) :rule trans :premises (t75 t44))
(step t77 (cl @p_12) :rule trans :premises (t70 t76))
(step t78 (cl @p_11) :rule not_or :premises (t58) :args (1))
(step t79 (cl false) :rule resolution :premises (t0 t77 t78) :args (@p_12 false @p_11 false))
(step t80 (cl (not false)) :rule false)
(step t81 (cl) :rule resolution :premises (t79 t80) :args (false true))

